rule_files:
  - ../rules/rules.yml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'fuse_triggered{circuit_id="1"}'
        values: '0 0 0 0 0 0 1 1 1 1'
    alert_rule_test:
      - eval_time: 10m
        alertname: FuseTriggered
        exp_alerts:
          - exp_labels:
              circuit_id: 1
            exp_annotations:
              summary: "Fuse triggered on circuit 1"
              description: Fuse triggered - a power grid is without power.

# No alert when fuse is not 1
  - interval: 1m
    input_series:
      - series: 'fuse_triggered{circuit_id="1"}'
        values: '0'
    alert_rule_test:
      - eval_time: 1m
        alertname: FuseTriggered
        exp_alerts:

# Alert when battery begins to drain
  - interval: 1m
    input_series:
      - series: 'battery_differential{circuit_id="1"}'
        values: '-1'
    alert_rule_test:
      - eval_time: 1m
        alertname: BatteryDraining
        exp_alerts:
          - exp_labels:
              circuit_id: 1
            exp_annotations:
              summary: "Battery draining on circuit 1"
              description: A battery has started to drain on circuit 1.
# No alert when we're charging
  - interval: 1m
    input_series:
      - series: 'battery_differential{circuit_id="1"}'
        values: '1'
    alert_rule_test:
      - eval_time: 1m
        alertname: BatteryDraining
        exp_alerts:

# Alert again when battery drained under 30 mins
  - interval: 1m
    input_series:
      - series: 'battery_seconds_empty{circuit_id="2"}'
        values: '3000 1500 1000 500'
      - series: 'battery_differential{circuit_id="2"}'
        values: '-1'
    alert_rule_test:
      - eval_time: 1m
        alertname: BatteryDrained
        exp_alerts:
          - exp_labels:
              circuit_id: 2
            exp_annotations:
              summary: "Under 30 minutes of battery left on circuit 2"
              description: A battery has nearly drained on circuit 2. Under thirty minutes of power remaining.

# No alert when charging
  - interval: 1m
    input_series:
      - series: 'battery_seconds_empty{circuit_id="1"}'
        values: '3700 1700'
      - series: 'battery_differential{circuit_id="1"}'
        values: '1 1'
    alert_rule_test:
      - eval_time: 10m
        alertname: BatteryDrained
        exp_alerts:
